<font color ='Navy'><center> __BruFace__ </center></font>
<br> MA1 Mechanical Engineering <br>  __Mechanical vibrations__ <br> <br> 
***
***
<font color='Navy'><h1><center>Session 3 : Multiple DOF system</center></h1></font> 
***
 <div class="alert alert-block alert-info">
<center><b>Welcome to the third session </b></center></div>
    
During the practical session, we will use the same libraries as the ones used in the previous sessions.

<div class="alert alert-block alert-info">
<center><b>Don't forget :It is important that you run each cell of the notebook. To do so select a cell (it will be highlighted) and press shift+enter on your keyboard or the play button in the menu above.</b></center></div>

```python
! pip install numpy 
! pip install scipy
! pip install matplotlib
! pip install ipywidgets
! pip install ipympl
```

```python
# let's import both numpy and matplotlib.pyplot 
# the "as" in the code bellow is to tell python that if we reference to the package using np and plt 
import numpy as np
import matplotlib.pyplot as plt 
import scipy.linalg as linalg
# don't worry about this line 
%matplotlib widget
```

# Exercises

## Exercise 1: Fundamental equations of a MDOF system

Consider the following three-degree-of-freedom (3 DOFs) system :
<img src="./Images/MDOF.png" width=200 height=200 />


<b> Assignment 1.1:</b> Write the equation of motion in the time domain for the system (in LaTeX)  in the cell below: 

(If you are not familiar with LaTeX you can use: https://www.hostmath.com/ and copy the LaTeX below, or write it on paper for yourself)

<!-- grade_id: cell-8688071c1ac5edbc, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

<b> Assignment 1.2:</b> Create a function `build_3DOF_matrices` that takes as arguments the stiffnesses $k$ and $k_1$, the damping values $c$ and $c_1$ and the mass $m$ and computes the system matrices ($M$, $C$, $K$). In the cell below we have given you a start, we show how you can use `np.array` to build the $M$-matrix. 

<!-- grade_id: cell-e1fffb4e22f46d21, points: 0, grade: false, solution: true -->
```python
def build_3DOF_matrices(m:float,c:float,c1:float,k:float,k1:float): 
    """
    Implement a function that computes the build the system matrix for a 3DOF matrix

    Arguments:
    m,k,k1,c,c1 -- parameters of the MDOF system

    Returns:
    M_mat : The M matrix from the system equation
    K_mat : The K matrix from the system equation
    C_mat : The C matrix from the system equation
  
    Tips use np.array([[],[],[]]) to define your 2D matrix
    check documentation FMI 'Numpy.array()'
    """  
    # here is an example of how to set a matrix 
    
    M_mat=np.array([[m, 0, 0],
           [0, m, 0],
           [0, 0, m]])
    
    # your code here
    raise NotImplementedError
    return M_mat, C_mat, K_mat
```

In the cell below we will check if you implemented the function correctly. To do so we will use several properties of the system matrices.  Do you understand all the checks that are being done?

<!-- grade_id: cell-868156988840d601, points: 1, grade: true, solution: false -->
```python
M, C, K = build_3DOF_matrices(m=1, k=10, k1=10, c=0.1, c1=0.1)

assert np.array_equal(M, np.eye(3)), 'M-matrix is not diagonal' # Mass matrix is a diagonal matrix

for matrix in [M,C,K]: # Iterate over the three matrices
    assert np.size(matrix,0) == np.size(matrix, 1), f'{matrix} is not square' # All three matrices should be square
    assert np.size(matrix,0) == 3, f'{matrix} is not 3x3' # For a 3DOF system we expect the system matrices to be 3X3 matrix
    assert np.allclose(matrix, matrix.T), f'{matrix} is not symmetric' # The matrices should be symmetric
    # The next one is a bit special, we assert whether or not the matrix is positive definite
    assert np.all(np.linalg.eigvals(matrix) > 0), f'{matrix} is not positive definite' 
# Some final checks
assert K[2,2] == 10
assert np.sum(K) == 10,  'The K matrix has an error in it'
assert np.sum(C) == 0.1, 'The C matrix has an error in it'

```

## Exercise 2 : Solving the conservative system

Before diving into a damped response we first will analyze the __conservative__ system response. Let us start with the eigenvalues and eigenvectors of the conservative system. To do so, we need to solve the following equation:

$(K-\lambda{}M)\Psi = (K-\omega^2M)\Psi =0$

Which is a basic eigenvalue problem. 

**Assignment 2.1:** In the cell below complete the function `calc_eigenvalues_eigenvectors` to solve the eigenvalue problem and return $\lambda$ and $\Psi$.

**TIP** : You will only need a single line of code, so have a look in the `scipy.linalg` module to find the right function.

<!-- grade_id: cell-b20f49cc2a2db1f5, points: 0, grade: false, solution: true -->
```python
def calc_eigenvalues_eigenvectors(M:np.array,K:np.array):
    """
    A function that computes the eigenvalues and eigenvectors
    based on the system parameter by reworking the equations of motion to a generalized eigenvalue problem

    Arguments:
    M,K -- system matrices of the MDOF system

    Returns:
    lamb -- The eigenvalues of the conservative eigenvalue problem
    psi -- The eigenvectors of the conservative eigenvalue problem
  
    
    """   
    # your code here
    raise NotImplementedError


    return lamb, psi
    
```

<b> Assignment 2.2:</b> Use `build_3DOF_matrices` and `calc_eigenvalues_eigenvectors`  to first calculate the eigenvalues `lamb` and eigenvectors `psi` (a.k.a the mode shapes).  Then calculate the eigenfrequencies in $Hz$ for a 3DOF system as presented above with $m = 1 kg$, $k = k1 = 16 N/m$, $c = c1 = 0.1 N s/m$.

Put your found eigenfrequencies in an array called `f_res`.

<!-- grade_id: cell-207549456677b63d, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

Let us have a look at the resonance frequencies you calculated:

<!-- grade_id: cell-97b06450290be31e, points: 2, grade: true, solution: false -->
```python
# Enumerate() method adds a counter to the list f_res and returns it in a form of elements_index and elements_values
for i,f_r in enumerate(f_res):
    print(f'Resonance frequency {i}: {np.real(f_r):0.2f} Hz')

# I do a little test, summing up the found resonance frequencies and comparing to the correct value
assert len(f_res) == 3, f'Your result contains {len(f_res)} resonance frequencies, we have 3 (pairs of) resonance frequencies'
assert round(np.sum(np.real(f_res)),2) != 12.73, 'Did you spot the power two that applies to omega?'
assert round(np.sum(np.real(f_res)),2) == 2.22, 'At least one of the resonance frequencies you found is incorrect!'
```

<b> Assignment 2.3:</b> In the field below answer following questions: 
* How many resonance frequencies did you obtain? Why?
* How many mode shapes did you obtain? Why?
* What are the dimensions of the mode shapes? Why?

<!-- grade_id: cell-ecb6aa6cbaf6f389, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

An important property of mode shapes is their orthogonality. In the next assignment you have to validate the orthogonality between two mode shapes.

Remember that :

$\Psi^TM\Psi=diag(\mu_i)$,

$\Psi^TK\Psi=diag(\mu_i\omega_i^2)$


<b> Assignment 2.4:</b> Complete `check_orthogonality` to calculate the orthogonality between two vectors. Consider using `np.matmul` and `np.transpose`: 

<!-- grade_id: cell-7179ce38f37e58df, points: 0, grade: false, solution: true -->
```python
def check_orthogonality(phi_1, phi_2, M):
    """
    A function that computes the orthogonality between two vectors

    Arguments:
    psi_1, psi_2 -- vectors of which the orthogonality is checked
    M    -- Mass matrix of the system

    Return: orth -- the orthogonality of the two vectors 
    
    TIP: matrix multiplications are done using np.matmul
    """   
    # your code here
    raise NotImplementedError

    return orth
```

Let us check whether your function does what it must do:

<!-- grade_id: cell-226a2ca25c40f171, points: 1, grade: true, solution: false -->
```python
# As a test let us take two orthogonal vectors
vec_1 = np.array([1,1,0])
vec_2 = np.array([1,-1,0])
M_test = np.eye(3) # The unity matrix

assert np.abs(check_orthogonality(vec_1, 5*vec_1, M_test)) > 1.e-10 # not equal to zero, a vector is not orthogonal to itself
assert np.isclose(check_orthogonality(vec_1, vec_2, M_test), 0) # = zero up to numerical precision
assert np.isclose(check_orthogonality(vec_1, 5*vec_2, M_test), 0) # = zero up to numerical precision
```

<b> Assignment 2.5:</b> Use `check_orthogonality` to check the orthogonality between your modeshapes:


<!-- grade_id: cell-415b0a5a73345f32, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

Were your mode shapes nicely orthogonal?

## Exercise 3 : The frequency response function

Let us continue with the 3DOF freedom model we used in Exercise 1. This time we will focus on the frequency response function. The frequency response function is a way to represent how a system will respond to a force being exerted on any of the degrees of freedom. For this we start with standard equations of motion of a MDOF system expressed in the frequency domain.

$\left(-\omega^2M+j\omega{}C+K\right)\vec{X}(\omega) = \vec{F}(\omega)$

Which we can rewrite as follows:

$\vec{X}(\omega) = \left(-\omega^2M+j\omega{}C+K\right)^{-1}\vec{F}(\omega) = H(\omega)\vec{F}(\omega)$

In which we call $H(\omega)$ the frequency response function.

Which we can rewrite as follows:

$\vec{X}(\omega) = \left(-\omega^2M+j\omega{}C+K\right)^{-1}\vec{F}(\omega) = H(\omega)\vec{F}(\omega)$

In which we call $H(\omega)$ the frequency response function.

In the cell below the function `frequency_response_function` calculates the FRF for a system defined using $M$, $C$ and $K$ matrices for a range of frequencies.

```python
def frequency_response_function(M:np.array, C:np.array, K:np.array, f= np.linspace(0,4,100), i=1,j=1):
    """
    A function that computes the Frequency Response Function from the system matrices

    Arguments:
    M,C,K -- system matrices of the MDOF system
    f -- the frequency for which the repsponse needs to be calculated
    i,j -- the specific element of the FRF matrix to be computed.

    Return: H_ij -- the i-th row and j-th column element of the FRF function at f
    """
    
    if not isinstance(f, np.ndarray):
        f = np.array([f])
    s = f*np.pi*2*1j
    
    M_s2 = np.outer(M,(s**2)).reshape(M.shape[0], -1, len(s))
    C_s = np.outer(C,(s**1)).reshape(M.shape[0], -1, len(s))
    K_s = np.outer(K,(s**0)).reshape(M.shape[0], -1, len(s))

    H_ij = []
    for s_i in range(len(s)):
        H_ij.append(np.linalg.inv((M_s2+C_s+K_s)[:,:,s_i])[i,j])
    
    return H_ij
```

<b> Assignment 3.1:</b> In the field below answer following questions:

* What are the dimensions of the FRF for a 3DOF system?
* Can you explain in your own words what the FRF represents?

<!-- grade_id: cell-af0c9be324cc2c95, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

<b> Assignment 3.2:</b> For the 3DOF system of exercise 1, can you calculate the transfer function between the input location of the force $f$ and the output location $x_3$ for a range of frequencies between 0 and 2Hz using `frequency_response_function`? Call your frequency vector `f` and your output `H_3f` and your system matrics `M`,`C` and `K`
    
**TIP**: Be aware that in python counting starts from 0, so the first column of a matrix is denoted 0



<!-- grade_id: cell-18dfc3b85db687b1, points: 0, grade: false, solution: true -->
```python
# your code here
raise NotImplementedError
```

Let us just quickly check if you have done the right steps.

<!-- grade_id: cell-54f3fab116aefadd, points: 2, grade: true, solution: false -->
```python
assert max(f) == 2 # Maximum frequency is 2Hz
assert len(H_3f) == len(f)




```

And let us have a look at your output:

```python
fig, axes = plt.subplots(2,1)

axes[0].semilogy(f,np.abs(H_3f))
axes[0].grid(which='both', linestyle=':')
axes[0].set_ylabel('Magnitude')
axes[1].plot(f,np.angle(H_3f,deg=True))
axes[1].set_ylabel('Angle (°)')
axes[1].set_xlabel('Frequency (Hz)')

axes[1].grid(which='both', linestyle=':')

t = axes[0].set_title('H_33(f)')
```

Obviously the properties of the FRF will change depending on the system properties. In the following interactive cell you can change the values of $m$, $c$, $k$. Can you identify the role of each of the parameters on the final FRF?

```python
import iplot.interactive_plot_1
```

As discussed in assignment 3.1 the FRF is in fact a 3 by 3 matrix for a three dimensional system. Up to this point we just focused on the transfer function between input location $f$ and output location $x_3(t)$. But how do the other elements of the transfer function look like?

<b> Assignment 3.3:</b> Update the code below to plot the FRF function for all 9 elements. Can you also plot a vertical line at the resonance frequencies on the x-axis? 
    
**TIP** : To plot a figure with logarithmic y-scale on a particular axis `ax` you can use the command `ax.semilogy` and to draw a single vertical line the same axis you can use the command `ax.vline(value)`.


<!-- grade_id: cell-21a07972bdd9e0e7, points: 0, grade: false, solution: true -->
```python
f = np.linspace(0,2,200)
n_dof = np.shape(M)[0]

fig, axes = plt.subplots(n_dof,n_dof)

for i in range(n_dof):
    for j in range(n_dof):
        # V---!ADD THE NECESSARY CODE HERE!---V
        # your code here
        raise NotImplementedError
        axes[i,j].set_title(f'$||H_{{{i+1}{j+1}}}||$')
        axes[i,j].set_xlabel('Frequency (Hz)')
        axes[i,j].grid(which='both', linestyle=':')
        axes[i,j].set_ylim([1e-3,1e1])
        axes[i,j].set_xlim([0, max(f)])
        
plt.subplots_adjust(left=0.1,
                    bottom=0.1, 
                    right=0.9, 
                    top=0.9, 
                    wspace=0.4, 
                    hspace=0.9)
```

<b> Assignment 3.4:</b> Can you answer following questions about the FRF's in the cell below?

* Do you notice any difference between the FRF's on the diagonal elements and those on non-diagonal elements?
* What do you notice about the resonance frequencies?
* What can you tell about the zeroes of the system?
* Can you spot other properties of the frequency response function?

<!-- grade_id: cell-2ae80da0bf5412b6, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

<div class="alert alert-block alert-success">
<b>That concludes this session:</b> In this session we tried to show you how MDOF systems can be modelled and how the FRF functions can be computed from the system matrices.
</div>
